```python
# Initialize Otter
import otter
grader = otter.Notebook("p1.ipynb")
```

```python
# PLEASE FILL IN THE DETAILS
# Enter None if you don't have a project partner

# project: p1
# submitter: lzhou256
# partner: None
# hours: ????
```

**Question 1:** "Hello, World!" program.

```python
name = "World"
greeting = "Hello, " + name + "!"
print(greeting)
```

Output:
```
Hello, World!
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:** Compute current year.

```python
curr_year = 45 * 45 - 3
print(curr_year)
```

Output:
```
2022
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

SUBMISSION INSTRUCTIONS: The zipfile automatically gets downloaded. You should have received GradeScope invite - please accept it. Login to GradeScope and upload the zipfile. Check otter results as soon as auto-grader gets completed.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False, run_tests=True)
```

Output:
```
Running your submission against local test cases...

File E:\Software\Anaconda3\lib\site-packages\otter\check\notebook.py:415, in Notebook.export(self, nb_path, export_path, pdf, filtering, pagebreaks, files, display_link, force_save, run_tests)
    413 if run_tests:
    414     print("Running your submission against local test cases...\n")
--> 415     results = grade_zip_file(zip_path, nb_path, self._path)
    416     print(
    417         "Your submission received the following results when run against " + \
    418         "available test cases:\n\n" + indent(results.summary(), "    "))
    420 if display_link:
    421     # create and display output HTML

File E:\Software\Anaconda3\lib\site-packages\otter\check\utils.py:86, in grade_zip_file(zip_path, nb_arcname, tests_dir)
     83     return results
     85 finally:
---> 86     os.remove(results_path)

PermissionError: [WinError 32] 另一个程序正在使用此文件，进程无法访问。: 'C:\\Users\\JselZ\\AppData\\Local\\Temp\\tmpyz3fshn3.pkl'
```

Tests for q1 (points: 12.5)
```
>>> greeting == "Hello, World!"
True
```


Tests for q2 (points: 12.5)
```
>>> curr_year == 2022
True
```
